---
title: 'Template Chapter 1'
description: 'This is a template chapter.'
---

## Ex 1.1

```yaml
type: NormalExercise
key: 84620e413f
lang: sql
xp: 100
skills: 1
```

Do some data science.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{sql}

```

`@solution`
```{sql}

```

`@sct`
```{python}

```

---

## Case Study: Data Driven Decision Making with SQL

```yaml
type: VideoExercise
key: d12207f973
xp: 50
```

`@projector_key`
dfa96b92ebbafae7cbc930ac36eff790
